# Big Data Pre-Processing

# Assignment 3

This notebook contains a set of exercises that will guide you through the different steps of this assignment. Solutions need to be code-based, i.e. hard-coded or manually computed results will not be accepted. Remember to write your solutions to each exercise in the dedicated cells. Do not modify the test cells. When you are done completing all the exercises submit this same notebook back to moodle in .ipynb format.

<div class="alert alert-success">The aim of this assignment is to download and inspect the Stop, Question and Frisk Data released by the New York Police Department every year. You will then use this dataset for visualization purposes during Session 7.</div>

<div class="alert alert-danger"><b>Submission deadline:</b> Sunday, November 7, 23:55</div>

## A bit of background

The stop-question-and-frisk program, or stop-and-frisk, in New York City, is a New York City Police Department practice of temporarily detaining, questioning, and at times searching civilians and suspects on the street for weapons and other contraband. This policy allowed police officers to stop, interrogate and search New York City citizens on the sole basis of “reasonable suspicion” for several years.

The stop-question-and-frisk program exploded into a national controversy during the mayoral administration of Michael Bloomberg, as the number of NYPD stops each year grew to hundreds of thousands. Most of the people stopped were Black and Latino, and nearly all were innocent. 

<img src='https://www.dropbox.com/s/uo6rj86eqvnusfy/stopandfrisk.png?raw=1' width=700>

In 1999, Blacks and Latinos made up 50 percent of New York’s population, but accounted for 84 percent of the city’s stops. Those statistics changed little in the following decade. Between 2004 and 2012, the New York Police Department made 4.4 million stops under the citywide policy. More than 80 percent of those stopped were Black and Latino people. The likelihood that a stopped African-American New Yorker yielded a weapon was half that of White New Yorkers stopped, and the likelihood of finding contraband on an African American who was stopped was one-third that of White New Yorkers stopped.

Stop-and-frisk peaked in 2011, when NYPD officers reported making nearly 700,000 stops.

In 2013, a federal judge found the New York City Police Department’s “Stop-and-Frisk” policy unconstitutional and ruled that NYPD’s stop-and-frisk tactics violate the U.S Constitution’s 4th Amendment prohibition of unreasonable searches and seizures.


## Retrieving the data

In this assignment you are going to download, inspect and prepare the data for their later use. Raw data records from the NYPD Stop, Question, and Frisk database are available for download at the following [link](https://www1.nyc.gov/site/nypd/stats/reports-analysis/stopfrisk.page). In this assignment you are going to use the data for the years 2003 through 2011 (both included). The corresponding databases are provided in separate zip archives.

Start by loading the data for 2003, the first year for which there are records.

<div class="alert alert-info">Write the code to read the data from the 2003 zip archive and store it in a DataFrame object called <i>df_2003</i>. Don't define any index when creating the DataFrame, simply revert to the default index values.</div>

<div class="alert alert-warning">Your code should import the data directly from their source location, i.e. should not refer to local files in your computer.</div>

<!-- grade_id: demo1, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
from io import BytesIO
from zipfile import ZipFile
from urllib.request import urlopen
url = urlopen("https://www1.nyc.gov/assets/nypd/downloads/zip/analysis_and_planning/stop-question-frisk/sqf-2003-csv.zip")

#Download Zipfile and create pandas DataFrame
zipfile = ZipFile(BytesIO(url.read()))
df_2003 = pd.read_csv(zipfile.open('2003.csv'), header=0, encoding='unicode_escape')
```

Output:
```
/Users/fabrizio/opt/anaconda3/lib/python3.8/site-packages/IPython/core/interactiveshell.py:3165: DtypeWarning: Columns (81,82,83,84,85) have mixed types.Specify dtype option on import or set low_memory=False.
  has_raised = await self.run_ast_nodes(code_ast.body, cell_name,
```

```python
pd.set_option('display.max_rows', 500)
pd.set_option('display.max_columns', 500)
pd.set_option('display.width', 1000)
```

```python
df_2003
```

Run the following cell to check your answer. If you get no error when running the cell it means that you did right. You can run this cell as many times as you want as long as you **don't modify it**.

<div class="alert alert-danger"> Make sure you pass this test successfully before moving on to the remaining exercises.</div>

<!-- grade_id: demo1_1, points: 0, grade: true, solution: false -->
```python
import unittest
import pandas as pd

# Variable df_2003 exists, is of type pandas.DataFrame, defaults to the standard index,
# has the right shape and includes the right column names
tc = unittest.TestCase()
tc.assertTrue('df_2003' in locals())
tc.assertIsInstance(df_2003, pd.core.frame.DataFrame)
tc.assertTrue(df_2003.index.equals(pd.core.indexes.range.RangeIndex(start=0, stop=160851, step=1)))
tc.assertEqual(df_2003.shape, (160851, 111))
tc.assertTrue(all(col in set(df_2003.columns)for col in ['dettypcm','ac_assoc','detailcm','ser_num','typeofid',
                                                         'explnstp','rf_verbl','sumissue','repcmd','ac_proxm',
                                                         'adtlrept','compyear','arstmade','ac_incid','cs_casng',
                                                         'city','othfeatr','contrabn','perstop','rf_rfcmp','searched',
                                                         'zip','timestop','pf_grnd','othpers','cs_cloth','rf_attir',
                                                         'sb_other','inout','crossst','perobs','pf_hands','cs_other',
                                                         'rf_knowl','officrid','cs_lkout','cs_furtv','comppct','sb_hdobj',
                                                         'cs_drgtr','rf_vcrim','ycoord','addrpct','post','sector','ac_other',
                                                         'offshld','xcoord','year','pf_drwep','ht_feet','pf_hcuff','riflshot',
                                                         'frisked','dob','pistol','pf_ptwep','premname','sb_outln','ht_inch',
                                                         'offunif','offverb','eyecolor','rf_othsw','rf_bulg','arstoffn',
                                                         'pf_baton','state','stname','sumoffen','beat','radio','age',
                                                         'addrnum','ac_time','datestop','cs_bulge','pct','cs_objcs','build',
                                                         'ac_evasv','crimsusp','othrweap','trhsloc','premtype','aptnum',
                                                         'pf_wall','linecm','stinter','ac_inves','weight','pf_pepsp','ac_stsnd',
                                                         'knifcuti','asltweap','addrtyp','haircolr','ac_cgdir','race',
                                                         'machgun','cs_descr','rf_furt','sex','sb_admis','rf_vcact','recstat',
                                                         'ac_rept','pf_other','cs_vcrim','revcmd','rescode']))
```

Now that you have successfully retrieved the data for one year, let's retrieve the data for the whole period 2003-2011.

<div class="alert alert-info"><b>Exercise 1 </b>Write the code to read the data from each zip archive and concatenate them all in a single DataFrame object called <i>df</i>. This DataFrame should contain all the data records from the NYPD Stop, Question, and Frisk database for the years 2003 through 2011. The oldest records should appear first.Don't modify the original column names. When you are done creating the DataFrame, reset the index so that it includes values from 0 to the length of df.<br><i>[1.5 points]</i></div>

<div class="alert alert-warning">As above, your code should import the data directly from their source location, i.e. should not refer to local files in your computer.</div>

<div class="alert alert-warning">When importing some of the files, you may get a warning stating that certain columns have mixed types. You can just ignore it. If it gets very annoying, set the <i>low_memory</i> parameter to False when reading the data. Note that as a consequence, it might take a bit longer to import all the data.</div>

<!-- grade_id: ex1, points: 0, grade: false, solution: true -->
```python
df = pd.DataFrame()
base = 'https://www1.nyc.gov/assets/nypd/downloads/zip/analysis_and_planning/stop-question-frisk/sqf-{}-csv.zip'
try:
  for i in range(2003,2012):
    url = urlopen(base.format(i))
    zipfile = ZipFile(BytesIO(url.read()))
    df_tmp = pd.read_csv(zipfile.open('{}.csv'.format(i)), header=0, encoding='ISO-8859-1', low_memory=False)
    df = pd.concat([df, df_tmp], verify_integrity=True, ignore_index=True)
except ValueError as e:
    print(e)
```

```python
df
```

The following cells contain the tests that will grade your code. **Don't modify them**.

<!-- grade_id: ex1_1, points: 0.1, grade: true, solution: false -->
```python
# LEAVE BLANK
```

<!-- grade_id: ex1_2, points: 0.75, grade: true, solution: false -->
```python
# LEAVE BLANK
```

<!-- grade_id: ex1_3, points: 0.25, grade: true, solution: false -->
```python
# LEAVE BLANK
```

<!-- grade_id: ex1_4, points: 0.4, grade: true, solution: false -->
```python
# LEAVE BLANK
```

<div class="alert alert-danger">This is the DataFrame that you are going to use throughout the assignment. Before moving on, check that it is properly created and named, that it includes all 123 columns and that it contains the data for all the years in the considered period, in the right order.</div>

Your DataFrame contains a very large number of records. In fact, it is so big that it may already be slowing down your computer. We are only interested in specific information. Hence, you are going to remove those columns that we are not going to use. To do so you first need to understand what each column name stands for.


## Understanding the data

Record layouts and code books for each year's data are recorded in separate Excel spreadsheet format files. A zip archive containing all years' spreadsheets can be downloaded from the NYPD website. In what follows, you are going to retrieve the code books from this archive and create a legend for your database.

Given that the data cover a 8-year period, it seems reasonable to assume that they have undergone some changes, i.e. that the columns are not exactly the same throughout the years. To check if this is the case, let's start by taking a look at the code books for 2003 and 2011, the first and last year in our considered period.

<div class="alert alert-info">Write the code to read the code book spreadsheets for the years 2003 and 2011 and store them in pandas DataFrame objects called <i>legend_2003</i> and <i>legend_2011</i>, respectively. These DataFrames should contain only two columns: <i>Variable</i> (the names of the columns) and <i>Label</i> (the descriptions for the different column names).</div>

<!-- grade_id: demo2, points: 0, grade: false, solution: true -->
```python
arch = 'https://www1.nyc.gov/assets/nypd/downloads/zip/analysis_and_planning/stop-question-frisk/SQF-File-Documentation.zip'
def zip_unarchiver (URL, year):
  url = urlopen(URL)
  zipfile = ZipFile(BytesIO(url.read()))
  df_exc = pd.read_excel(zipfile.open('{} SQF File Spec.xlsx'.format(year)), na_values='NA', usecols=[0,2], index_col=None)
  df_exc = df_exc.iloc[3:]
  df_exc = df_exc.reset_index(drop=True)
  if year >= 2008:
    var_def = 'NYPD Stop Question Frisk Database  {}'.format(year)
  else:
    var_def = 'NYPD Stop Question Frisk Database {}'.format(year)
  df_exc.rename(columns={var_def: 'Variable', 'Unnamed: 2': 'Label'}, inplace=True)
  return df_exc
```

```python
legend_2003 = zip_unarchiver(arch, 2003)
legend_2011 = zip_unarchiver(arch, 2011)
```

The lengths of both DataFrames are different. This means that there are slight variations in the columns for each year. As a result, you will need to import the code books for all the different years to create your legend and ensure that all column names are properly documented.

<div class="alert alert-info"><b>Exercise 2 </b>Write the code to read the code book spreadsheets for all the years in the considered period and store the <b>unique values</b> in a single DataFrame called <i>legend</i>. This DataFrame should contain only two columns: <i>Variable</i> (the names of the columns) and <i>Label</i> (the descriptions for the different column names).<br><i>[1.5 points]</i></div>

<div class="alert alert-warning">Once again, your code should import the data directly from their source location, i.e. should not refer to local files in your computer.</div>

<div class="alert alert-warning">Make sure that your <i>legend</i> DataFrame contains all the required entries and no duplicates.</div>

<!-- grade_id: ex2, points: 0, grade: false, solution: true -->
```python
legend = pd.DataFrame()
try:
  for i in range(2003,2012):
    tmp = zip_unarchiver(arch, i)
    legend = pd.concat([legend, tmp], verify_integrity=True, ignore_index=True).drop_duplicates()
    legend = legend.reset_index(drop=True)
except ValueError as e:
    print(e)
```

```python
legend
```

The following cells contain the tests that will grade your code. **Don't modify them**.

<!-- grade_id: ex2_1, points: 0.1, grade: true, solution: false -->
```python
# LEAVE BLANK
```

<!-- grade_id: ex2_2, points: 0.5, grade: true, solution: false -->
```python
# LEAVE BLANK
```

<!-- grade_id: ex2_3, points: 0.5, grade: true, solution: false -->
```python
# LEAVE BLANK
```

<!-- grade_id: ex2_4, points: 0.4, grade: true, solution: false -->
```python
# LEAVE BLANK
```

If you were able to successfully retrieve the information from all the code books, you'll notice that the columns included in the documentation differ from those in the *df* DataFrame you created above: all the column names included in *legend* are also included in *df*, but the opposite is not true. This is because there are several spelling mistakes in the databases you imported in Exercise 1.

<div class="alert alert-info"><b>Exercise 3 </b>Write the code to identify the columns from <i>df</i> that are not included in the code books. Store their names in a list called <i>typos</i>.<br><i>[0.25 points]</i></div>

<!-- grade_id: ex3, points: 0, grade: false, solution: true -->
```python
typos = []
for column in df.columns:
  if column not in list(legend['Variable']):
    typos.append(column)
#typos = list(set(df.columns).difference(set(legend['Variable'].unique())))
```

```python
typos
```

The following cells contain the tests that will grade your code. **Don't modify them**.

<!-- grade_id: ex3_1, points: 0.1, grade: true, solution: false -->
```python
# LEAVE BLANK
```

<!-- grade_id: ex3_2, points: 0.15, grade: true, solution: false -->
```python
# LEAVE BLANK
```

Dedicate a few minutes to revise the list of columns in <i>legend</i>. Identify the correct column names for each of the entries in <i>typos</i>. Notice that the differences are mostly due to spelling mistakes, or uppercase vs. lowercase issues. There's one single case where the correspondence is not clear: <i>'wepfound'</i>. You may need to check the documentation once again to identify the right column name for this value. 

<div class="alert alert-info"><b>Exercise 4 </b>Write the code to modify the column names in <i>df</i> so that those column names included in <i>typos</i> are all renamed to match those in <i>legend</i>. Once you are done renaming the columns, merge the columns with the same names to eliminate the duplicates.<br><i>[2 points]</i></div>

<div class="alert alert-warning">Check that all column names in <i>df</i> are written exactly as in the code books and that there are no duplicated columns after this exercise.</div>

<!-- grade_id: ex4, points: 0, grade: false, solution: true -->
```python
corrected_names = ["dettypCM","lineCM", "detailCM","othrweap", "dettypCM", "detailCM", "detailCM", "rescode", "premtype", "premname", "addrnum", "stname", "stinter", "addrpct"] 
XDict= dict(zip(typos, corrected_names))
df.rename(columns = XDict, inplace= True)
```

```python
df = df.groupby(by=df.columns, level=0, axis=1).first()
```

```python
#!python -m pip install 'fsspec>=0.3.3'
#import dask.dataframe as dd
#df_reduced = dd.from_pandas(df, 2)
#result = df_reduced.groupby(by=df.columns,axis=1).first().compute() parallelized result
#result = df_reduced.sum(level=0,axis=1, skipna=True).compute()
#df.groupby(by=df.columns, level=0, axis=1).first()
#result = df.groupby('id').max().reset_index().compute()
#result = df.groupby('id').max().reset_index().compute()
```

The following cells contain the tests that will grade your code. **Don't modify them**.

<!-- grade_id: ex4_1, points: 0.5, grade: true, solution: false -->
```python
# LEAVE BLANK
```

<!-- grade_id: ex4_2, points: 0.5, grade: true, solution: false -->
```python
# LEAVE BLANK
```

<!-- grade_id: ex4_3, points: 0.5, grade: true, solution: false -->
```python
# LEAVE BLANK
```

<!-- grade_id: ex4_4, points: 0.5, grade: true, solution: false -->
```python
# LEAVE BLANK
```

Now that you have cleaned the column names and merged the data, let's remove all unnecesary information. The list shown below contains the descriptions for the variables we are interesting in.

```python
information = ["YEAR OF STOP (CCYY)", "PRECINCT OF STOP (FROM 1 TO 123)", "DATE OF STOP (MM-DD-YYYY)", 
               "CRIME SUSPECTED", "WAS AN ARREST MADE ?", "PERIOD OF STOP (MMM)",
               "WAS A SUMMONS ISSUED ?", "WAS A PISTOL FOUND ON SUSPECT ?", "WAS A RIFLE FOUND ON SUSPECT ?", 
               "WAS AN ASSAULT WEAPON FOUND ON SUSPECT ?", "WAS A KNIFE OR CUTTING INSTRUMENT FOUND ON SUSPECT ?",
               "WAS A MACHINE GUN FOUND ON SUSPECT ?","WAS ANOTHER TYPE OF WEAPON FOUND ON SUSPECT", 
               "PHYSICAL FORCE USED BY OFFICER - HANDS", "PHYSICAL FORCE USED BY OFFICER - SUSPECT AGAINST WALL",
               "PHYSICAL FORCE USED BY OFFICER - SUSPECT ON GROUND", "PHYSICAL FORCE USED BY OFFICER - WEAPON DRAWN",
               "PHYSICAL FORCE USED BY OFFICER - WEAPON POINTED", "PHYSICAL FORCE USED BY OFFICER - BATON",
               "PHYSICAL FORCE USED BY OFFICER - HANDCUFFS", "PHYSICAL FORCE USED BY OFFICER - PEPPER SPRAY",
               "PHYSICAL FORCE USED BY OFFICER - OTHER", "SUSPECT'S RACE", "SUSPECT'S SEX"]
```

<div class="alert alert-info"><b>Exercise 5 </b>Write the code to identify the columns that contain the information above and remove all the other columns from your <i>df</i> DataFrame. Don't modify neither the name of the Dataframe nor any of the given column names when doing so.<br><i>[1 points]</i></div>

<!-- grade_id: ex5, points: 0, grade: false, solution: true -->
```python
newcols = []
for info in information:
    newcols.append(legend.loc[legend['Label'] == info, 'Variable'].iloc[0])
df = df[newcols]
```

The following cells contain the tests that will grade your code. **Don't modify them**.

<!-- grade_id: ex5_1, points: 0.5, grade: true, solution: false -->
```python
# LEAVE BLANK
```

<!-- grade_id: ex5_2, points: 0.5, grade: true, solution: false -->
```python
# LEAVE BLANK
```

## Enriching the data

The column *pct* encodes the precinct where each stop took place. Police units based in these precincts are responsible for safety and law enforcement within a designated geographic area. There are a total of 77 precincts in the city of New York, each identified with a number from 1 to 123. These precincts can be further aggregated into the 5 major boroughs of New York: Bronx, Brooklyn, Manhattan, Queens and Staten Island. Let's enrich the data by adding the borough data so that we can later visualize it using different levels of granularity.

To do so, you will need to look for the correspondences between boroughs and police precincts, and add this information to your database. Before doing so, however, you will need to make sure that the *pct* column is properly encoded.

<div class="alert alert-info"><b>Exercise 6 </b>Write the code to standardize the values in the <i>pct</i> column by converting all entries to integers. Don't modify the name of the column when doing so and don't add any new columns to the DataFrame either.<br><i>[0.5 points]</i></div>

<div class="alert alert-warning">There is at least one value in the original <i>pct</i> column which cannot be converted to an integer. You can simply drop the corresponding entries from the DataFrame</div>

<!-- grade_id: ex6, points: 0, grade: false, solution: true -->
```python
df['pct'] = pd.to_numeric(df['pct'], errors = 'coerce')
df.dropna(subset = ['pct'], inplace = True)
df['pct'] = df['pct'].astype(int)
```

```python
df['pct']
```

The following cells contain the tests that will grade your code. **Don't modify them**.

<!-- grade_id: ex6_1, points: 0.1, grade: true, solution: false -->
```python
# LEAVE BLANK
```

<!-- grade_id: ex6_2, points: 0.1, grade: true, solution: false -->
```python
# LEAVE BLANK
```

<!-- grade_id: ex6_3, points: 0.1, grade: true, solution: false -->
```python
# LEAVE BLANK
```

<!-- grade_id: ex6_4, points: 0.1, grade: true, solution: false -->
```python
# LEAVE BLANK
```

<!-- grade_id: ex6_5, points: 0.1, grade: true, solution: false -->
```python
# LEAVE BLANK
```

Now that the *pct* column is properly encoded, let's include the information regarding the borough.

<div class="alert alert-info"><b>Exercise 7 </b>Write the code to create a new column called <i>borough</i>. This column should contain the name of the borough that corresponds to the precinct of each entry in your <i>df</i> DataFrame. Consider only the 5 major boroughs in New York City, as written above, and write their full names. By the end of this exercise only the <i>borough</i> column should be added to <i>df</i>.<br><i>[1.5 point]</i></div>

<div class="alert alert-warning">There some entries in the original DataFrame for which the values in the <i>pct</i> column are not properly informed. These rows should be dropped at this step, since we can't assign them a borough.</div>

<!-- grade_id: ex7, points: 0, grade: false, solution: true -->
```python
def get_borough_pct(column):
    pct = column["pct"]
    if pct < 35:
        return "Manhattan"
    elif pct >=40 and pct <= 52:
        return "Bronx"
    elif pct >= 60 and pct <= 94:
        return "Brooklyn"
    elif pct >=100 and pct <= 115:
        return "Queens"
    elif pct >=120 and pct<=123:
        return "Staten Island"
    else:
        return None
```

```python
df['borough'] = df.apply(get_borough_pct,axis=1)
```

The following cells contain the tests that will grade your code. **Don't modify them**.

<!-- grade_id: ex7_1, points: 0.1, grade: true, solution: false -->
```python
# LEAVE BLANK
```

<!-- grade_id: ex7_2, points: 0.25, grade: true, solution: false -->
```python
# LEAVE BLANK
```

<!-- grade_id: ex7_3, points: 0.15, grade: true, solution: false -->
```python
# LEAVE BLANK
```

<!-- grade_id: ex7_4, points: 0.25, grade: true, solution: false -->
```python
# LEAVE BLANK
```

<!-- grade_id: ex7_5, points: 0.75, grade: true, solution: false -->
```python
# LEAVE BLANK
```

Let's further filter down these data. Some of the columns store related information. For example, lots of the columns refer to *PHYSICAL FORCE USED BY OFFICER*. Since we are not interested in the specific form of force used, but rather on the existence of such force, let's aggregate the information for these columns.

The same thing can be said for the set of columns that contain the phrase WAS A - FOUND ON SUSPECT ?. Once again, we are not interested in the actual type of weapon found, but on whether a weapon was indeed found on the suspect. Hence, let's aggregate all these data into a single column.

Finally, let's also aggregate the data about whether an arrest was made. We'll consider that an arrest took place if an arrest was actually made or if a summons was issued, or both.

<div class="alert alert-info"><b>Exercise 8 </b>Write the code to create new columns called <i>physical_force</i>, <i>weapon</i> and <i>arrest</i>. These columns should take the value 1 in <b>int</b> form if any of the related columns are encoded as "Y" and 0 in <b>int</b> form otherwise.<br><i>[0.75 points]</i></div>

<!-- grade_id: ex8, points: 0, grade: false, solution: true -->
```python
pf_cols = ['pf_hands', 'pf_wall', 'pf_grnd', 'pf_drwep', 'pf_ptwep', 'pf_baton', 'pf_pepsp', 'pf_hcuff', 'pf_other']
weap_cols = ['pistol', 'riflshot', 'asltweap', 'knifcuti', 'machgun', 'othrweap' ]
arr_cols = ['arstmade', 'sumissue']
```

```python
def get_pf_cols(column):
    for colm in pf_cols:
        pf = column[colm]
        if 'Y' in pf:
            return int(1)
        else:
            return int(0)
```

```python
def get_weap_cols(column):
    for colm in weap_cols:
        pf = column[colm]
        if 'Y' in pf:
            return int(1)
        else:
            return int(0)
```

```python
def get_arr_cols(column):
    for colm in arr_cols:
        pf = column[colm]
        if 'Y' in pf:
            return int(1)
        else:
            return int(0)
```

```python
df['physical_force'] = df.apply(get_pf_cols,axis=1)
```

```python
df['physical_force'] = df['physical_force'].astype(int)
```

```python
df['weapon'] = df.apply(get_weap_cols,axis=1)
```

```python
df['arrest'] = df.apply(get_arr_cols,axis=1)
```

The following cells contain the tests that will grade your code. **Don't modify them**.

<!-- grade_id: ex8_1, points: 0.1, grade: true, solution: false -->
```python
# LEAVE BLANK
```

<!-- grade_id: ex8_2, points: 0.15, grade: true, solution: false -->
```python
# LEAVE BLANK
```

<!-- grade_id: ex8_3, points: 0.5, grade: true, solution: false -->
```python
# LEAVE BLANK
```

We are now going to focus on the <i>race</i> column. This column encodes critical information for our use case. You can find more information regarding the different letter codes in the documentation provided by NYPD. We will begin by changing the name of the column to <i>ethnicity</i>. Ideally, we would prefer to keep the original values for this variable. However, for visualization purposes, we will need to group the values into larger categories. 

<div class="alert alert-info"><b>Exercise 9 </b>Write the code to rename the <i>race</i> column as <i>ethnicity</i>. Then write the code to modify the values of your <i>ethnicity</i> column. This column should take the value 'Black' if the original record was encoded as Black or Black-Hispanic. It should take the value 'White' if the original record was encoded as White or White-Hispanic. It should take the value 'Asian' if the original record was encoded as Asian/Pacific Islander. It should take the value 'Other' if the original record was encoded as American Indian/Alaskan native, Other, Unknown or Not listed.<br><i>[0.75 points]</i></div>

<div class="alert alert-warning">Make sure to check that by the end of this exercise your DataFrame contains no new columns. You can add them to help you, but make sure to remove them before moving on.</div>

<!-- grade_id: ex9, points: 0, grade: false, solution: true -->
```python
df.rename(columns = {'race': 'ethnicity'}, inplace = True)
```

```python
black = ['B', 'P']
white = ['W', 'Q']
asian = ['A']
other = ['I', 'X', 'Z', 'NOT LISTED', 'U']
```

```python
for i, value in enumerate(df['ethnicity']):
    if value in black:
        df.at[i,'ethnicity'] = 'Black'
    
    elif value in white:
        df.at[i,'ethnicity'] = 'White'
        
    elif value in asian:
        df.at[i,'ethnicity'] = 'Asian'
        
    elif value in other:
        df.at[i,'ethnicity'] = 'Other'
```

The following cells contain the tests that will grade your code. **Don't modify them**.

<!-- grade_id: ex9_1, points: 0.1, grade: true, solution: false -->
```python
# LEAVE BLANK
```

<!-- grade_id: ex9_2, points: 0.15, grade: true, solution: false -->
```python
# LEAVE BLANK
```

<!-- grade_id: ex9_3, points: 0.5, grade: true, solution: false -->
```python
# LEAVE BLANK
```

Finally, let's remove redundant information.

<div class="alert alert-info"><b>Exercise 10 </b>Write the code to remove all columns except <i>year</i>, <i>pct</i>, <i>datestop</i>, <i>perstop</i>, <i>crimsusp</i>, <i>ethnicity</i>, <i>sex</i>, <i>borough</i>, <i>physical_force</i>, <i>weapon</i>, and <i>arrest</i> from your DataFrame <i>df</i>.<br><i>[0.25 points]</i></div>

<div class="alert alert-warning">Make sure to double check that you remove all the columns above. At the end of this exercise your DataFrame should contain the columns you selected in exercise 2 plus the 4 new columns you created in exercises 4, 5, and 6, minus those listed above.</div>

<!-- grade_id: ex10, points: 0, grade: false, solution: true -->
```python
columns_new = ['year', 'pct', 'datestop', 'perstop', 'crimsusp', 'ethnicity', 'sex', 'borough', 'physical_force', 'weapon', 'arrest']
df = df[columns_new]
df
```

The following cells contain the tests that will grade your code. **Don't modify them**.

<!-- grade_id: ex10_1, points: 0.1, grade: true, solution: false -->
```python
# LEAVE BLANK
```

<!-- grade_id: ex10_2, points: 0.15, grade: true, solution: false -->
```python
# LEAVE BLANK
```

Congrats! You are now finished preparing the data. As always, you can also check the exercises below for extra points.

<div class="alert alert-danger"><b>Bonus 1 </b>Write the code to retrieve the datetime values for every entry in your DataFrame <i>df</i>. Don't modify any column names when doing so. Don't add any new columns either. Make sure that by the end of this exercise the <i>datestop</i> column is of type <i>datetime64[ns]</i>.<br><i>[1 point]</i></div>

<!-- grade_id: b1, points: 0, grade: false, solution: true -->
```python
for i, date in enumerate(df['datestop']):
    try:
        date = pd.to_datetime(date, format='%m%d%Y').strftime("%d/%m/%Y")
        df['datestop'].iloc[i] = date
    except:
        try:
            date = pd.to_datetime(date, format='%Y-%m-%d').strftime("%d/%m/%Y")
            df['datestop'].iloc[i] = date
        except:            
            try:
                date = pd.to_datetime(date, format='%d%m%Y').strftime("%d/%m/%Y")
                df['datestop'].iloc[i] = date
            except:
                pass


#turn into datetime64[ns]
df['datestop'] = df['datestop'].astype('datetime64[ns]')
df['datestop']
```

The following cells contain the tests that will grade your code. **Don't modify them**.

<!-- grade_id: b1_1, points: 0.1, grade: true, solution: false -->
```python
# LEAVE BLANK
```

<!-- grade_id: b1_2, points: 0.4, grade: true, solution: false -->
```python
# LEAVE BLANK
```

<!-- grade_id: b1_3, points: 0.5, grade: true, solution: false -->
```python
# LEAVE BLANK
```

<div class="alert alert-danger"><b>Bonus 2 </b>Write the code to count the number of entries corresponding to males for the different categories of the <i>ethnicity</i> columns in your DataFrame <i>df</i>. Store the resulting values in a dictionary called <i>distribution</i>, where the keys should correspond to the different values of the <i>ethnicity</i> column and the values to the frequency of male records for each.<br><i>[0.5 points]</i></div>

<!-- grade_id: b2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
pivot = pd.pivot_table(df, index = ['ethnicity'], columns = 'sex', values = 'arrest', aggfunc = 'count')['M']
df_dist = pd.DataFrame(pivot)
df_dist
```

```python
df_dist['M'] = df_dist['M'].astype(int)
df_dist
distribution = df_dist.to_dict()['M']
distribution
```

The following cells contain the tests that will grade your code. **Don't modify them**.

<!-- grade_id: b2_1, points: 0.1, grade: true, solution: false -->
```python
# LEAVE BLANK
```

<!-- grade_id: b2_2, points: 0.4, grade: true, solution: false -->
```python
# LEAVE BLANK
```